# Molecular Modelling Exercises

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

---

# Jupyter Notebooks

First a quick introduction to the general usage of Jupyter Notebook

```python
# This is a code cell, that you can execute
# Try:
# Enter (to edit the cell)
# Ctrl-Enter (to execute without moving)
# Shift-Enter (to run and move to the next cell)
# Alt-Enter (to run and insert a new cell after this cell)
# Escape    (to deselect the cell and be able to move around)
# All the variables defined previously will be available to the other cells.
# Run this cell
a = 1
```

Then continue with running the next cells:

```python
print(a)
```

Output:
```
1
```

Now run this one twice and then the one above:

```python
a = a + 1
```

The code in all cells is in the same scope, so if you change a variable and re-run an earlier cell, then also in that earlier cell you are working with the changed variable!

# Python

Over the course of the next 4 Notebooks, you will learn about the programming language Python.
If something is unclear please don't hesitate and ask the lecturers and tutors.
If you want further information, [the python documentation](https://docs.python.org/3/) covers more tutorials and topics.


Python, unlike most other programming languages, is not statically typed and compiled but interpreted. 
In Python you don't have to specify what type of value you assign to a variable, and you can re-assign a value of different type to a variable.

```python
a = 'this is a string'
print(type(a))
a = 1
print(type(a))
```

Output:
```
<class 'str'>
<class 'int'>
```

You can use Python just like one would use a calculator.
Feel free to explore:

```python
1 + 2
```

Output:
```
3
```

```python
# feel free to explore further
```

If you want to store the result of a calculation, you can assign it to a variable:

```python
variable_name = expression
```

You can name your variables, almost, however you want.
The Python Interpreter does not care how you name your variables, or how long your variable names are. Variable names are for humans to remember what values are stored in that variable!

Variable names can be composed of letters, underscores and numbers, but must not start with a number.
In general, it is convention to use lowercase letters and seperate words with underscores (called 'snake_case').

```python
addition = 3 + 2
subtraction = 3 - 2
```

To view the value of an variable you can `print` the variable using the `print`-function: 

```python
print('The result of the addition is: ')
print(addition)
```

Output:
```
The result of the addition is: 
5
```

Also, Jupyter Notebooks display the return value of the last expression.

```python
print('The result of the subtraction is not printed but rather the return value of variable name')
subtraction
```

Output:
```
The result of the subtraction is not printed but rather the return value of variable name
1
```

## Comments

Comments are ignored by the interpreter. They are useful if you want to write a short explanation for your future self what that part of code is doing.

Everything after a `#` is ignored by the interpreter.

```python
# the + 3 below is ignored, as well as this line here
1 + 2 # + 3
```

Output:
```
3
```

## Exercise 1

Store any value you want in a variable called `exercise_variable`
The cell with `assert` tests wether you solved the problem. There may be hidden tests.

<!-- grade_id: cell-9446f34ee163b2e2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
exercise_variable = 27 + 3 / 5
```

<!-- grade_id: cell-3f1a04248e83f061, points: 1, grade: true, solution: false -->
```python
assert 'exercise_variable' in locals()
```

# Indentation

In python, indentation of specific code blocks is important. While in other languages, curly brackets and semicolons are used to denote code blocks that belong together.
To indent code, 4 spaces (in many editors you can use 1 tab for this) should be used.

```python

```