# Molecular Modelling Exercises

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

---

# Datatypes

# Numeric Datatypes

Integers and Floats are two common numeric types of python with which you can perform calculations.
In Python you do not have to explicitly convert inbetween ints and floats for calculation.

## Integer
The name origins from Latin meaning: _in_ ("not") _tangere_ ("to touch"). As the name suggests, integers are 'whole' numbers without a floating point.
Integers are exact. For large numbers you can add `_` inbetween numbers as you want for better readability.


```python
integer_number = 3
print('The type of integer_number is: ', type(integer_number))

large_number = 1_000_000_000_000 # for better readability you can add underscores
print('The underscores are not shown and only serve for readability purposes:', large_number)
```

Output:
```
The type of integer_number is:  <class 'int'>
The underscores are not shown and only serve for readability purposes: 1000000000000
```

## Floats

Floats on the other hand can store rational numbers. The disadvantage is, that computers cannot calculate exactly with floating point numbers. Usually 7-14 digits are trustworthy, depending on the precision used.

```python
float_number = 3.1415
print(type(float_number))

float_number = 2.
print(type(float_number))
```

Output:
```
<class 'float'>
<class 'float'>
```

```python
# showcase rounding errors
print('.1+.1+.1=', .1+.1+.1, ' and not 0.3')
.1 + .1 + .1 == .3
```

Output:
```
.1+.1+.1= 0.30000000000000004  and not 0.3
False
```

You can perform arithmetic operations between ints and floats without explicitly converting the data types.

```python
division = integer_number / float_number 
multiplication = integer_number * float_number
exponentiation = integer_number ** float_number
```

## Exercise: Double-Bond-Equivalents

write a function that takes the number of hydrogen, carbon, nitrogen and halogen atoms and calculates the double bond equivalent:
\begin{equation}
    DBE = C - \frac{H - N + X}{2} + 1
\end{equation}


<!-- grade_id: cell-5ad9cec94c7a56ac, points: 0, grade: false, solution: true -->
```python
def double_bond_equivalent(carbons, hydrogens, nitrogens, halogens):
    # This is a function here, you will learn more about it later.
    # make sure to create a variable `dbe` that will store the result of the calculation
    # also make sure that all of the code is indented to the same level as the return dbe statement
    
    # YOUR CODE HERE
    dbe = carbons - ((hydrogens - nitrogens + halogens)/2)+1
    return dbe
```

<!-- grade_id: cell-53e647fa1fa05863, points: 1, grade: true, solution: false -->
```python
# this code down below tests wether your function works correctly

assert double_bond_equivalent(6, 6, 0, 0) == 4
assert double_bond_equivalent(1, 2, 1, 1) == 1

```

# Strings

Strings in Python store text information. They are enclosed in quotation marks.

```python
conventional_string = 'Conventionally, strings are enclosed in single quotation marks in python.'
also_accepted_string = "But many programmers are used to use double quotation marks."

longer_multi_line_string = '''This is why you can use both single and double quotation marks.
If you want to write longer strings you need to use three (single or double) quotation marks 
    to start
and also three quotation marks to indicate 
    the end 
of the string.'''
```

```python
print(conventional_string)
print(also_accepted_string)
print(longer_multi_line_string)
```

Output:
```
Conventionally, strings are enclosed in single quotation marks in python.
But many programmers are used to use double quotation marks.
This is why you can use both single and double quotation marks.
If you want to write longer strings you need to use three (single or double) quotation marks 
    to start
and also three quotation marks to indicate 
    the end 
of the string.
```

As you saw in the multiline string, even the linebreaks and indentation in the string is preserved.

## Basic string formatting

Strings have many methods implemented for them, which allows you to format them as you want.

```python
hello = 'Hello '
world = 'World'

print(hello.lower(), world.upper(), hello.capitalize(), world.swapcase())
```

Output:
```
hello  WORLD Hello  wORLD
```

Sometimes, it you want to add the value of variables into a string. One way to do this is using string concatenation (as you will see below). 

Another approach is using the format feature of python strings. Take a look at the example below.

The method `format` is called on the string. It replaces the curly brackets `{}` with the value you give to the `format` method.
The second example uses the `f-strings`, which is a shorthand notation for the same functionality.

Store your name in the `name` variable below.

<!-- grade_id: cell-9952f1709a7a87a1, points: 0, grade: false, solution: true -->
```python
# enter your name below
name = "Michael Fill"
# YOUR CODE HERE

greeting = 'Hello, {}'.format(name.title())
greeting2 = f'Hello, {name.title()}'

print(greeting)

greeting == greeting2
```

Output:
```
Hello, Michael Fill
True
```

The arithmetic operators that you saw for the numerical data types also exist for strings. But here they do something different.

```python

hello + world
```

Output:
```
'Hello World'
```

```python
hello * 3
```

Output:
```
'Hello Hello Hello '
```

If you read in data from a text file, all the information will be stored as string which can lead to suprising results if you then do calculations without further ado.

```python
a = '6'
b = '3.1415'

print('a + b=', a+b)
print('a * 3=', a*3)
```

Output:
```
a + b= 63.1415
a * 3= 666
```

In these cases, you need to convert your strings to the proper numerical type.

```python
a = int(a)
b = float(b)

print('a + b=', a+b)
print('a * 3=', a*3)
```

Output:
```
a + b= 9.1415
a * 3= 18
```

# Boolean Values
Boolean values can be either `True` or `False`. The result of comparison operators are boolean values.
The operators to check for:

- equality is `==`
- inequality is `!=`
- bigger / smaller are `>` / `<`
- bigger or equal / smaller or equal are `>=` / `<=`

```python
a = 1 
b = 1
c = 2

first_condition = a == b
second_condition = a >= c

print('first_condition = ', first_condition)
print('second_condition =', second_condition)
```

Output:
```
first_condition =  True
second_condition = False
```

With Boolean values you can perform logical operations:
The three basic logical operations are:
 - conjunction - in python as `and`
 - disjunction - in python as `or`
 - negation    - in python as `not`

```python
print('and:', first_condition and second_condition)
print('or:', first_condition or second_condition)
print('not:', not first_condition)
```

Output:
```
and: False
or: True
not: False
```

# Exercise: Boolean Values

Add a line in both the even as well as the uneven function so that the function returns a boolean value, that determines wether the number provided is even or uneven.

hint: use the `%` modulo operator that returns the remainder of an integer division
```python
>>> 5 % 3
2
```

<!-- grade_id: cell-d9c6a994a436236e, points: 0, grade: false, solution: true -->
```python
def even(number):
    # YOUR CODE HERE
    expr = number % 2 == 0
    return expr
```

<!-- grade_id: cell-270723ee4a267515, points: 0, grade: false, solution: true -->
```python
def uneven(number):
    # YOUR CODE HERE
    expr = number % 2 == 1
    return expr
```

<!-- grade_id: cell-765f6108254ab9f3, points: 1, grade: true, solution: false -->
```python
assert even(2) == True
assert even(1) == False

assert uneven(1) == True
assert uneven(2) == False
```

```python

```