# Molecular Modelling Exercises

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

---

# Collections

Usually one resorts to programming languages if one wants to process a lot of data. To store lots of data, Python provides different types of collections that can store as much data as you can fit in your memory.

You already met one type of collections: ```String```.
Strings are also collections that store text. So everything you learn here, can be applied to Strings as well.

# Lists and Tuples

Lists and Tuples are the most basic types of collections that exist.
The main difference between lists and tuples is, that tuples can not be changed after creation, while lists can be dynamically changed, extended and shortened.

Lists are indicated by square brackets `[]` while tuples are indicated by round brackets `()`.


```python
numbers_to_ten = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

lat_lon_innsbruck = (47.2692, 11.4041)
```

In python you can store different types of data in the same list or tuple. You can also store other lists or variables in a list.

In the example below, we store a tuple containing a string, a float and an integer in a list.

```python
# [name, weight, atomic number]
oxygen = ('Oxygen', 15.999, 8)
elements = [oxygen]
```

To access a specific element of a list or a tuple (or even string) you can indicate which element you want with square brackets after the variable name. Be aware that Python is 0 indexed, meaning, the first element has index 0, not 1! `variable_name[index]`

If you want to access a slice of a collection, meaning a range of elements, you can also access that with square brackets. This time, instead of a single index for a single element one provides a start and a stop elements, where the stop element is not included. Start and stop indices are seperated by a colon: `variable_name[start:stop]`.

![indexing in python](https://cdn.programiz.com/sites/tutorial2program/files/python-list-index.png)

```python
print(oxygen[-1])
print(numbers_to_ten[8])
print(numbers_to_ten[1:8])
```

Output:
```
8
9
[2, 3, 4, 5, 6, 7, 8]
```

Lets take a closer look at lists:
To add a further element to a list you can `append` it to the list.

```python
hydrogen = ('Hydrogen', 1.008, 1)
helium = ('Helium', 4.002, 2 )
elements.append(hydrogen)
elements.append(helium)

elements
```

Output:
```
[('Oxygen', 15.999, 8), ('Hydrogen', 1.008, 1), ('Helium', 4.002, 2)]
```

To remove the first occurance of an element you can `remove` it.

```python
elements.remove(helium)
elements
```

Output:
```
[('Oxygen', 15.999, 8), ('Hydrogen', 1.008, 1)]
```

In lists, as they are dynamic, you can assign new values as well.

```python
print(elements[0])
elements[0] = hydrogen

elements
```

Output:
```
('Oxygen', 15.999, 8)
[('Hydrogen', 1.008, 1), ('Hydrogen', 1.008, 1)]
```

In tuples, removing, appending and assigning new values to a element does not work because tuples are immutable (=cannot be changed).

If you want to quickly get the length of a collection, one can use the `len` function.

```python
len(elements)
```

Output:
```
2
```

# Dictionary

Sometimes a more complex collection type might be required. Dictionaries allow the mapping of one value to another. 
Similar as before, an element of a dictionary can be accessed by providing the key in square brackets after the dictionary name. In contrast to lists, the 'index' now is no longer the position in the list, but the key you provided when creating the dictionary. The keys must be unique.

If the key does not exist in the dictionary, Python will complain and throw an error.

```python
# the indentation here in the example below is only for readability and won't influence the code.

elements = {'hydrogen' : (1, 1.008),
            'helium' : (2, 4.002), 
            'oxygen' : (8, 15.999),
            }

elements
```

Output:
```
{'hydrogen': (1, 1.008), 'helium': (2, 4.002), 'oxygen': (8, 15.999)}
```

```python
elements['hydrogen']
```

Output:
```
(1, 1.008)
```

Similar to lists and tuples, the keys and associated values can be of any kind. The only requirement for the key is, that it must not be of a data type that can change (such as list).

```python
elements['random_key'] = 'This will add a new entry into the dictionary with key "random_key" and this string as value.'

print(elements)

elements['random_key'] = 'You can also change the associated value for a key.'
```

Output:
```
{'hydrogen': (1, 1.008), 'helium': (2, 4.002), 'oxygen': (8, 15.999), 'random_key': 'This will add a new entry into the dictionary with key "random_key" and this string as value.'}
```

One can very quickly check wether a dictionary contains a certain key.

```python
'hydrogen' in elements
```

Output:
```
True
```

Again there are a lot of different methods (functions on objects) for dictionaries, so we won't cover all of them here and just showcase one.
`pop` returns the associated value and removes this entry from the dictionary.

```python
removed_entry = elements.pop('random_key')
print('removed_entry =', removed_entry)
print(elements)
```

Output:
```
removed_entry = You can also change the associated value for a key.
{'hydrogen': (1, 1.008), 'helium': (2, 4.002), 'oxygen': (8, 15.999)}
```

# Control Flow

Control Flow plays a central part in a program or script and controls in what order or which code is executed at a time.

## If and else

As the name suggests, with `if` and `else` you can control which code block your program executes based on a condition.
This is the first example where the indentation starts to matter.

```python
if 'hydrogen' in elements: # this is True
    print('Hydrogen is in elements.')
else:
    print('It\'s not.')
print('Code at this indentation level is not affected by the if/else ')
```

Output:
```
Hydrogen is in elements.
Code at this indentation level is not affected by the if/else
```

```python
if 'carbon' in elements: # This is False
    print('An entry for carbon exists.')
else:
    print('There is no carbon in the dictionary.')
print('Code at this indentation level is not affected by the if/else ')
```

Output:
```
There is no carbon in the dictionary.
Code at this indentation level is not affected by the if/else
```

if you want to distinguish more than 2 cases, you can denote the cases inbetween the first and last with `elif`
```python
if first_condition and second_condition:
    # do something
    ...
elif first_condition or second_condition:
    # do something else
    ...
elif not first_condition:
    # do something based on this condition
    ...
else:
    # none of the above apply
    ...
```


## For loops
If you want to do an operation for every element in a collection, for loops are the way to go.
Again, everything that should be done for every element in a collection, has to be indented once below the `for`.

### with range
`range` is a function that returns values from the start (normally 0 if not provided) up until the endpoint.

```python
numbers_to_hundred = [] # empty list
for number in range(1, 101): # go from 1 to 100, the 101 is a result from the 0 indexing
    numbers_to_hundred.append(number) # add the number to the list from line 1

numbers_to_hundred[-1] # display the last element
```

Output:
```
100
```

```python
sum_of_all_squares = 0 
for i in range(1_000_000): # iterate up until 1 million
    sum_of_all_squares += i  # short hand notation for sum_of_all_squares = sum_of_all_squares + i

sum_of_all_squares
```

Output:
```
499999500000
```

`sum_of_all_squares = sum_of_all_squares + i` is not a valid mathematical function, it should be False, right?

The `=` operator must not be confused for the mathematical `=`, but rather is an assignment operator.
The 'mathematical' `=` corresponds to the `==` operator in python. 
In the above example the value stored in `sum_of_all_squares + i` is assigned to the variable name `sum_of_all_squares`.
This is repeated for all `i` from 0 to 1 million.

Note that 1 million is excluded in the range above!

### over collections

In python you can iterate over lists and tuples directly.

```python
usernames = ['alice', 'bob', 'foo', 'bar', 'baz']

print('-'*80)
# the most `pythonic way` is to iterate directly over the elements of the list
for user in usernames:
    print(user)

print('-'*80)
# but you can also iterate over the list using indeces
for i in range(len(usernames)): # len() is a function that returns the length of a sequence
    print(usernames[i])

print('-'*80)
# Alternatively the more pythonic way, if you need the index number, 
# but also want to iterate over the list directly you can use the 
# enumerate function.
for i, user in enumerate(usernames):
    print(user, i, usernames[i])
```

Output:
```
--------------------------------------------------------------------------------
alice
bob
foo
bar
baz
--------------------------------------------------------------------------------
alice
bob
foo
bar
baz
--------------------------------------------------------------------------------
alice 0 alice
bob 1 bob
foo 2 foo
bar 3 bar
baz 4 baz
```

If you iterate over dictionaries directly, you cannot use the range(len(...)) approach, this works this way.

```python
elements = {1: 'hydrogen', 2: 'helium', 3:'berrylium'}

print('-'*80)
for atomic_number, name in elements.items():
    print(f'The atomic number is: {atomic_number}, the associated name is: {name}')


print('-'*80)
# if you only need the keys
for atomic_number in elements: # or elements.keys(), both are equivalent for iteration.
    print(atomic_number)

print('-'*80)
# if you only need the values
for name in elements.values():
    print(name)
```

Output:
```
--------------------------------------------------------------------------------
The atomic number is: 1, the associated name is: hydrogen
The atomic number is: 2, the associated name is: helium
The atomic number is: 3, the associated name is: berrylium
--------------------------------------------------------------------------------
1
2
3
--------------------------------------------------------------------------------
hydrogen
helium
berrylium
```

## Exercise 1: lists and for loops

Create a list called `exercise_1` that contains all elements from 10 to 120 (including 10 and including 120).

<!-- grade_id: cell-97369c7603873a67, points: 0, grade: false, solution: true -->
```python
exercise_1 = []
# YOUR CODE HERE
for i in range (10,121):
    exercise_1.append(i)
```

<!-- grade_id: cell-fad48b799477a983, points: 1, grade: true, solution: false -->
```python
assert exercise_1[0] == 10
assert exercise_1[-1] == 120
assert len(exercise_1) == 111
```

# Functions and Methods

You won't need to write your own python functions, but you should be able to use them. Functions in python can be used as follows:
```python
function(argument, ..)
```

To define your own functions the syntax looks like this:
```python
def function(argument, ..):
    # insert code that should be executed
    value = ...
    # if you want to return a value, this is done like this
    return value
```

Depending on the function, different numbers of arguments or datatypes might be required. Note again, that everything that is indented below the function definition is part of the function.

Below are some examples of built-in python functions.

```python
l = [8, 1, 2, 1, 5]
```

```python
max(l) # returns the highest element of a sequence
```

Output:
```
8
```

```python
len(l)
```

Output:
```
5
```

```python
abs(-8) # returns the absolute value
```

Output:
```
8
```

```python
sum_to_hundred = sum(numbers_to_hundred) # returns the sum of all values in a sequence
sum_to_hundred
```

Output:
```
5050
```

In python everything is an object. Many objects have so-called methods associated with them.
You already saw one of them. For example the `list_name.append()` is a method that directly operates on the list and is only available for objects of type list.
The usage of methods is equivalent to that of a regular function.

```python
# note the difference between the sort method and the sorted function
print(f'Initial list= {l}')
print(f'{sorted(l) }')
print(f'after function call= {l}')
l.sort()
print(f'after method call= {l}')
```

Output:
```
Initial list= [8, 1, 2, 1, 5]
[1, 1, 2, 5, 8]
after function call= [8, 1, 2, 1, 5]
after method call= [1, 1, 2, 5, 8]
```

## Exercise 2

What is the differnece between the ```list.sort()``` method and the ```sorted()``` function?


<!-- grade_id: cell-5fe2ed4553ae3ed0, points: 1, grade: true, solution: true -->
list.sort() calls the sort method provided by any list object and actually changes the list to its now sorted state. the sorted() function called on a list returns the sorted list, but does not change the list object.

## Exercise 3: if statements

Complete the following function, so that the function assigns a grade based on the points that a student received for a test.

The `grade` should be assigned based on these criterions:
- 1: above 87.5% 
- 2: 75-87.5%
- 3: 62.5-75%
- 4: 50-62.5%
- 5: below 50%

<!-- grade_id: cell-38ed15d7bdea055d, points: 0, grade: false, solution: true -->
```python
def calculate_grade(achieved_points, total_points=24):
    if achieved_points/total_points*100 >= 87.5:
        grade = 1
    elif achieved_points/total_points*100 >= 75:
        grade = 2
    elif achieved_points/total_points*100 >= 62.5:
        grade = 3
    elif achieved_points/total_points*100 >= 50:
        grade = 4
    else:
        grade = 5
    return grade
```

<!-- grade_id: cell-0fd3c612fa84376c, points: 1, grade: true, solution: false -->
```python
assert calculate_grade(24, 24) == 1
assert calculate_grade(87.5, 100) == 1
assert calculate_grade(9, 10) == 1
assert calculate_grade(8, 10) == 2
assert calculate_grade(5, 10) == 4
```

```python

```